Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "GISRSArtists"
COLLABORATORS = "Anne-Juul Welsink and Marrit Leenstra"
```

---

# Assignment

Wageningen University wants to determine the percentage of water on some parcels at the campus. Can you help?

## Determine how much area (m2 and percentage) of the combined WUR parcels are covered by water by using the BGT WFS (typename = bgt:waterdeel) and visualize both the water area and parcels in a map.

Wageningen University are specifically interested in parcels with "perceelnummer": ['10709', '10905', '10906', '10907', '10908', '10909', '11184', '11185', '11208', '11593']

Work with a project structure. There is a jupyter notebook file (.ipynb) made available for you.
You must fill in your code and assign your answers to specific names as specified within. You can import functions from a scripts folder as usual. If you find working in spyder more comfortable, you may first develop your solution in spyder, then copy over your code and check if it runs.

In your script, after downloading and selecting correct parcels, make a bounding box of the parcels which you will use to download water areas (hint .envelope, .exterior). You can create these functions:


    geocodePlacenameToCoordinate
    reproject
    downloadWFSToGeoDataFrame
    calculatePercentageArea
    makeWebmap or makeMap (for visualization you can choose to make a map with Folium or matplotlib)
    getBoundingBox


### assign your answers to variables with the name 'answer_fraction' (a fraction of 1) and 'answer_m2'

Go easy on webservices! There is no need to use for loops or while statements to call web services multiple times! Upload your documented and well structured Python scripts to a GitLab repository.

Hint: If you need some help how to make functions or how to import modules, have a look at the code in the Python refresher or in the cheat sheets mentioned in the Python refresher.



```python
#space for imports, feel free to add your own or add another cell with imports in them:
import json
from owslib.wfs import WebFeatureService
import geopandas as gpd
import pandas as pd
from geopy.geocoders import Nominatim
from pyproj import Proj, transform
from python.reproject import reproject
from python.geocodePlacenameToCoordinates import geocodePlacenameToCoordinates
from python.downloadWFSToGeoDataFrame import downloadWFSToGeoDataFrame
```

<!-- grade_id: cell-a71c1037b0ca57da, points: 0, grade: false, solution: true -->
```python
# get coordinates of Wageningen
try:
    wageningen = geocodePlacenameToCoordinates('Wageningen University')
except:
    print("Location not found")

# reproject to RD New
x, y = reproject(wageningen, '4326', '28992')
print(lng, lat)

# download WFS parcels
parcelsWFSUrl = 'https://geodata.nationaalgeoregister.nl/kadastralekaartv3/wfs'
waterWFSUrl = 'https://geodata.nationaalgeoregister.nl/beta/bgt/wfs'
version = '2.0.0'
parcelsBBox = x-450, y+350, x-300, y+350
parcelGDF = downloadWFSToGeoDataFrame(parcelsWFSUrl, version, parcelsBBox, 'kadastralekaartv3:perceel')
print(parcelGDF['perceelnummer'])


# select parcels
#parcelArray = ['545', '320', '10906', '10907', '10908', '10909', '11184', '11185', '11208', '11593']
#parcelArray = ['10709', '10905', '10906', '10907', '10908', '10909', '11184', '11185', '11208', '11593']
print(parcelGDF[parcelGDF['perceelnummer'].isin(['10906', '10907', '10908', '10909', '11184', '11185', '11208', '11593'])])

#print(selectedParcelsGDF)


```

Output:
```
173994.1578792833 444133.60329471016
/home/WUR/leens006/miniconda3/envs/vectorpython/lib/python3.7/site-packages/owslib/feature/wfs200.py:120: FutureWarning: The behavior of this method will change in future versions. Use specific 'len(elem)' or 'elem is not None' test instead.
  if serviceproviderelem:
0    10708
1     8782
2    10384
3     7233
Name: perceelnummer, dtype: object
Empty GeoDataFrame
Columns: [geometry, kadastraleGemeenteCode, kadastralegrootte, logischtijdstipOntstaan, logischtijdstipVervallen, lokaalID, perceelnummer, perceelnummer_deltax, perceelnummer_deltay, perceelnummer_rotatie, plaatscoordinaten, sectie]
Index: []
```

```python
print(parcelGDF['perceelnummer'].isin(['10709', '10905', '10906', '10907', '10908', '10909', '11184', '11185', '11208', '11593']))

```

Output:
```
0      False
1      False
2      False
3      False
4      False
5      False
6      False
7      False
8      False
9      False
10     False
11     False
12     False
13     False
14     False
15     False
16     False
17     False
18     False
19     False
20     False
21     False
22     False
23     False
24     False
25     False
26     False
27     False
28     False
29     False
       ...  
970    False
971    False
972    False
973    False
974    False
975    False
976    False
977    False
978    False
979    False
980    False
981    False
982    False
983    False
984    False
985    False
986    False
987    False
988    False
989    False
990    False
991    False
992    False
993    False
994    False
995    False
996    False
997    False
998    False
999    False
Name: perceelnummer, Length: 1000, dtype: bool
```

<!-- grade_id: cell-7b28aeb052277bbb, points: 6, grade: true, solution: false -->
```python
#This cell is here to help you with finding trivial mistakes.

assert type(answer_fraction) == float  , 'check if answer_fraction is a number with decimal'
assert (answer_fraction<1.0) == True , 'check if answer_fraction is below 1'
assert type(answer_m2) == float , 'check if answer_m2 is a number with decimal'

```

<!-- grade_id: cell-7ac7b1827f71a075, points: 4, grade: true, solution: false -->
```python
#This cell has hidden tests in it, don't remove or move. It won't return any errors for you.

```

### Now plot a map of the water features and the administrative parcels surrounding them

<!-- grade_id: cell-37b9d31dd6747386, points: 1, grade: true, solution: true -->
```python
#Plot the map

# YOUR CODE HERE
raise NotImplementedError()
```

### Bonus: Calculate how much area (m2 and fraction) is water, per parcel. Present result nicely summarized in one table and also in a nice-looking map visualization (further below) of the relevant parcels and their (partial) coverage by water.

#### to be awarded points, assign your answer to a pandas DataFrame object named 'answer_bonus', which in turn accepts a python dictionary as an input. Note that the "column names" (called 'keys' when in a dictionary) have to be parcelID, m2, and fraction (order does not matter)

example: 

`answer_dict = {'parcelID': [1, 2], 'm2': [3, 4], 'fraction': [5, 6]}`

`answer_bonus=pandas.DataFrame(data=answer_dict)`

to show your table, just type line with answer_bonus and hit enter, or use print(answer_bonus)

<!-- grade_id: cell-71f4e5f5395d85d4, points: 0, grade: false, solution: true -->
```python
#Section for calculation and present table

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-790384f59e4b5666, points: 1, grade: true, solution: false -->
```python

assert type(answer_bonus) == pd.core.frame.DataFrame, 'check if its a pandas dataframe'
assert 'parcelID' and 'm2' and 'fraction' in list(answer_bonus.columns.values) , 'check if the columns are named correctly'

```

<!-- grade_id: cell-9e0cd5287f7cb064, points: 1, grade: true, solution: true -->
```python
#Section for plotting map

# YOUR CODE HERE
raise NotImplementedError()
```